# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

### 21.1 What does the keyword `def` mean before a function name?

> A) _define_ - You are saying that this is the function definition
>
> B) _default_ - This is the default work the function should do
>
> C) _def_ - Doesn't stand for anything, it's just what you put there
>
> D) _defer_ - Meaning don't do anything right now

```python
answer211 = "A"
```

### 21.2 What is the minimum number of parameters that you have to provide for this function to work?

```python
def my_function (fname='John', lname='Doe', gender='M'):
    ...
```

> A) 3
>
> B) 2
>
> C) 1
>
> D) 0

```python
answer212 = "D"
```

### 21.3 Which of the following lines of code does not have any functions in it?


> A) `value = math.log(23,2)`
>
> B) `pi * (radius ** 2)`
>
> C) `[ x for x in range(10) ]`
>
> D) `print('I am a good student')`

```python
answer213 = "B"
```

### 22.1 Which is true about docstrings?

> A) They are stored in Python as just a long string inside the function definition.
>
> B) They can contain tests to help you demonstrate your code works correctly.
>
> C) They tell Python what to say when someone asks for `help()` on your function.
>
> D) All of these

```python
answer222 = "D"
```

### 22.2 True or False: From now on you should include docstrings in the _applied_ homework assignments?

Type `answer222 = True` or `answer222 = False` as your response

```python
answer222 = True
```

### 23.1 Which of the following are typical ways of encapsulating logical rules or mapping in Python?

> A) dict
>
> B) str
>
> C) function
>
> D) A and C
>
> E) None of these

```python
answer231 = "C"
```

### 23.2 What does DRY stand for?


> A) Do not repeat yourself
>
> B) Do not repeat yourself
>
> C) Do not repeat yourself
>
> D) All of the above

```python
answer232 = "D"
```

### 24.1 When Python tells you there's been an exception, what is the key word that we use in our Python code to indicate how Python should react to that exception?

> A) Throw
>
> B) Raise
>
> C) Except
>
> D) Crash

```python
answer241 = "C"
```

### 24.2 When you, as the programmer, want to notify Python that an exception should be triggered, what is the key word we use for that?

> A) Throw
>
> B) Raise
>
> C) Except
>
> D) Crash

```python
answer242 = "B"
```

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(L):
    """(list) -> int
    This code will accept a list of stings and return as output the position number of the longest string. 
    >>> longest(['Someday', 'I', 'will', 'be', 'able', 'to', 'dance']) 
    0
    >>> longest(['I', 'am', 'watching', 'the', 'BattleHawks'])
    4
    >>> strings = ['happy', 'birthday', 'to', 'me']
    >>> longest(strings)
    1
    """
    position = 0
    big_num = 0
    Num_lt_str = 0
    n=0
    for string in L:
        for i in range(len(string)):
            Num_lt_str = Num_lt_str + 1
        if position == 0 and big_num == 0:
            big_num = Num_lt_str
            position = position
        elif big_num < Num_lt_str:
            big_num = Num_lt_str 
            position = n
        else:
            big_num = big_num 
            position = position
        Num_lt_str = 0
        n = n + 1
    return position
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['Someday', 'I', 'will', 'be', 'able', 'to', 'dance']) 
Expecting:
    0
ok
Trying:
    longest(['I', 'am', 'watching', 'the', 'BattleHawks'])
Expecting:
    4
ok
Trying:
    strings = ['happy', 'birthday', 'to', 'me']
Expecting nothing
ok
Trying:
    longest(strings)
Expecting:
    1
ok
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

### 25.2 Exception Handling

Write a function called `safe_average(L)` to calculate the average of a list of values.  If any of the values in the list are not valid numbers, then the function should simply ignore them and exclude from the average.  If none of the values are valid numbers, then the function should return None.

Example:

```
>>> safe_average([100,200,'three'])
150

>>> safe_average(['one','two','three'])
None
```

<!-- grade_id: exception-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def safe_average(L):
    """(list) -> int
    Returns the average from a list of numbers.  It will only report the averages of integers.  If there are other types of information, the data will be excluded from the calculation.  
    >>> safe_average([100,200,200.1])
    150
    >>> safe_average([100,200,'three'])
    150
    >>> safe_average([145,'money', 100000, 234.5, "fun", 354])
    33499
    >>> safe_average(['one','two','three'])
    None
    """
    n = 0
    total = 0
    ave = 0
    for i in L:
        if type(i) is int:
            n = n + 1
            total = total + i
        else:
            n = n
            total = total
    try:
        ave = total/n
        ave = int(ave)
        return ave
    except ZeroDivisionError:
        print (None)

### END SOLUTION
```

```python
doctest.run_docstring_examples(safe_average, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    safe_average([100,200,200.1])
Expecting:
    150
ok
Trying:
    safe_average([100,200,'three'])
Expecting:
    150
ok
Trying:
    safe_average([145,'money', 100000, 234.5, "fun", 354])
Expecting:
    33499
ok
Trying:
    safe_average(['one','two','three'])
Expecting:
    None
ok
```

<!-- grade_id: exception-tests, points: 10, grade: true, solution: false -->
```python
assert safe_average([1,2,3]) == 2
assert safe_average([1,'two',3]) == 2
assert safe_average(['one','two','three']) == None
```

Output:
```
None
```

---

## Rubric for 25.1 and 25.2

* Create a properly working function
* Follow good practices for naming the function and its parameters
* Provide meaningful documentation
* Provide sufficient docstring tests