# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape #there are 12000 rows with 39 columns
display(hhcaps.columns)
hhcaps.head()
```

Output:
```
Index(['State', 'CMS Certification Number (CCN)*', 'Provider Name', 'Address',
       'City', 'Zip', 'Phone', 'Type of Ownership',
       'Offers Nursing Care Services', 'Offers Physical Therapy Services',
       'Offers Occupational Therapy Services',
       'Offers Speech Pathology Services', 'Offers Medical Social Services',
       'Offers Home Health Aide Services', 'Date Certified',
       'HHCAHPS Survey Summary Star Rating',
       'Footnote HHCAHPS Survey Summary Star Rating',
       'Star Rating for health team gave care in a professional way',
       'Footnote for Star Rating for gave care in a professional way',
       'Percent of patients who reported that their home health team gave care in a professional way',
       'Footnote for percent of patients who reported that their home health team gave care in a professional way',
       'Star Rating for health team communicated well with them',
       'Footnote for Star Rating for communicated well with them',
       'Percent o
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = 0
answer = len(hhcaps.columns) #gives the length of the columns list
print(answer)

rows, columns = hhcaps.shape #this is another way of shoing the rows and columns
print(rows)
print(columns)
## Solution goes here
```

Output:
```
39
12000
39
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

```python

```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = 0
display(hhcaps.head(5))
answer = len(hhcaps['State'].unique())
print(answer)
## Solution goes here
```

Output:
```
State  CMS Certification Number (CCN)*  \
0    AL                            17000   
1    AL                            17008   
2    AL                            17009   
3    AL                            17013   
4    AL                            17014   

                                       Provider Name  \
0  BUREAU OF HOME & COMMUNITY SERVICES           ...   
1  JEFFERSON COUNTY HOME CARE                    ...   
2  ALACARE HOME HEALTH & HOSPICE                 ...   
3  GENTIVA HEALTH SERVICES                       ...   
4  AMEDISYS HOME HEALTH                          ...   

                                             Address                  City  \
0  201 MONROE STREET, THE RSA TOWER,  SUITE 1200 ...  MONTGOMERY             
1  2201 ARLINGTON AVENUE                         ...  BESSEMER               
2  2970 LORNA ROAD                               ...  BIRMINGHAM             
3  1239 RUCKER BLVD                              ...  ENTERPRISE             
4  68278 M
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

```python

```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
answer = ''
answer = hhcaps['State'].\
    value_counts().\
    index[0]
#looking at the state column
#counting the values for each
#listing the top value in the index
answer

## Solution goes here
```

Output:
```
'TX'
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

```python

```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
answer = ''
x  = hhcaps.groupby('State') #grouping by state
answer = x['Star Rating for health team communicated well with them'].\
    agg(['mean']).\
    sort_values(by='mean', ascending = False).\
    index[0]
#looking at the Star Rating columns and aggrecating the mean then sorting the values from highest and displaying the first in the index
## Solution goes here

```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

```python

```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
import numpy
## Solution goes here
answer = 0.0
display(hhcaps['Type of Ownership'].unique())

answer = hhcaps[hhcaps['Type of Ownership'] == 'Hospital Based Program']['Star Rating for how patients rated overall care from agency'].\
    agg(['mean'])
answer = round(answer,5) #showing the rounded answer
display(answer.dtype) #displaying the type.  The assert lines didn't work so I shoed that the answer and types were correct
answer

```

Output:
```
array(['Official Health Agency', 'Local', 'Visiting Nurse Association',
       'Hospital Based Program', 'Combination Government Voluntary',
       'Skilled Nursing Facility Based Program',
       'Rehabilitation Facility Based Program'], dtype=object)
dtype('float64')
mean    3.69252
Name: Star Rating for how patients rated overall care from agency, dtype: float64
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
# import numpy
# assert(type(answer) == float or type(answer) == numpy.float64)
# assert(round(answer,5) == round(3.6925207756232687,5))

```

```python

```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

```python
answer = 'test'
display(hhcaps.shape)
c = len(hhcaps.columns)
special = hhcaps.nunique() #lists the number of unique entries per column
n=0
display(special[0])
special
for i in range(c):              #using a for loop in the range of the number of columns to show the column that has only 1 unique value
    if special[i] == 1 and n == 0:
        n = n+1
        answer = special.index[i]
    elif special[i] == 1 and n > 0:  #I put this line in place for the situation that there was more than one cloumn that had the same value in every row
        answer1 =  special.index[i]
        answer = answer + ', ' + answer1
answer
```

Output:
```
(12000, 39)
55
'Offers Nursing Care Services'
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())


```

```python

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
## Solution goes here
import re
answer = 0
adjust = hhcaps
column = hhcaps.columns #listing the column names
r, c = adjust.shape     #counting the row and columns
i = 0

for col in column:                          #using a for loop the match all the columns with offers and removing all other columns
    match = re.match(r"^Offers.*", col)
    if (match):
        adjust
        z = match
    else:
        adjust = adjust.drop([col], axis = 1)
for i in range(r):                         #using a for loop to show add when all of the values are the same
    w = adjust.loc[i].value_counts()       #and also check that the value is True
    if w[1] == 6  and adjust[adjust.columns[1]].loc[i] == True: 
        answer = answer+1

answer

```

Output:
```
9029
```

```python
display(adjust.columns[1])

```

Output:
```
'Offers Physical Therapy Services'
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

```python

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
answer = 0
answer = hhcaps['Number of completed Surveys'].\
    replace('Not Available', 'NAN').\
    median()
#looks at the Number of completed Surveys column, replaces the not available with NAN the calculated the median
## Solution goes here
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

```python

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

```python
answer = 0

stl= (hhcaps['City'].str.strip() == 'SAINT LOUIS').\
    value_counts()                                    #strips the column city of extra values to the right and left looks at for Saint louis only
display(stl)                                          #and counts how many are true and false
answer = int(stl[1])                                  #answer is given for all of the true values
answer
```

Output:
```
False    11977
True        23
Name: City, dtype: int64
23
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

```python

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
answer = 0
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'

state = hhcaps[['State', field]]                         #looks at just the 2 columns state and field

state = state[state[field] != 'Not Available']           #separates all of the rows that do not have not available

state[field] = state[field].astype(str).astype(int)      #takes the vaules and turns them into integers

states = state.groupby('State').mean()                   #groups the values by state and calculates the mean
answer = states.sort_values(field, ascending=False).index[0] #answer is sorted from highest to lowest and listing only the top in the index


# ## Solution goes here

```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```

```python

```